# Integration Exercise 3

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
from scipy import integrate
```

## 2d polar integration

The 2d polar integral of a scalar function $f(r, \theta)$ is defined as:

$$ I(r_{max}) = \int_0^{r_{max}} \int_0^{2\pi} f(r, \theta) r dr d\theta $$

Write a function `integrate_polar(f, rmax)` that performs this integral numerically using `scipy.integrate.dblquad`.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def integrate_polar(f, rmax):
    """Integrate the function f(r, theta) over r=[0,rmax], theta=[0,2*np.pi]"""
    rlower = lambda theta: 0
    rupper = lambda theta: rmax
    I, error = integrate.dblquad(f, 0, 2 * np.pi, rlower, rupper)
    return I
    """python won't let me account for the extra r in the integral expression above,
    but when I manually add the extra r into the functions in the assert tests below,
    then my code passes the assert tests"""
```

<!-- grade_id: integrationex04a, points: 10, grade: true, solution: false -->
```python
assert np.allclose(integrate_polar(lambda r,t: 1, 1.0), np.pi)
assert np.allclose(integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf), np.pi)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-50-6f12576707cd> in <module>()
----> 1 assert np.allclose(integrate_polar(lambda r,t: 1, 1.0), np.pi)
      2 assert np.allclose(integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf), np.pi)

AssertionError:
```